rule_files:
- rules.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: zookeeper_status_code{project_name="default"}
    values: "5x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Is_Degraded_Alarm
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        project_name: default
        service: zookeeper-montemplate
      exp_annotations:
        description: ZooKeeper is Degraded.
        summary: Some of ZooKeeper Service pods are down

- interval: 1m
  input_series:
  - series: zookeeper_status_code{project_name="default"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Is_Degraded_Alarm
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: zookeeper_status_code{project_name="default"}
    values: "10x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Is_Down_Alarm
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        project_name: default
        service: zookeeper-montemplate
      exp_annotations:
        description: ZooKeeper is Down.
        summary: All of ZooKeeper Service pods are down

- interval: 1m
  input_series:
  - series: zookeeper_status_code{project_name="default"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Is_Down_Alarm
    exp_alerts: []

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", pod="zookeeper-0"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits_cpu_cores{exported_namespace="default", exported_pod="zookeeper-0"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_CPU_Load_Alarm
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: zookeeper-montemplate
      exp_annotations:
        description: ZooKeeper CPU load is higher than 95 percents
        summary: Some of ZooKeeper Service pod loads CPU higher then 95 percents

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", pod="zookeeper-0"}
    values: "0x5"
  - series: kube_pod_container_resource_limits_cpu_cores{exported_namespace="default", exported_pod="zookeeper-0"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_CPU_Load_Alarm
    exp_alerts: []

- interval: 1m
  input_series:
  - series: container_memory_working_set_bytes{namespace="default",container="zookeeper-0",pod="zookeeper-0"}
    values: "1x5"
  - series: kube_pod_container_resource_limits_memory_bytes{exported_namespace="default", exported_pod="zookeeper-0"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Memory_Usage_Alarm
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: zookeeper-montemplate
      exp_annotations:
        description: ZooKeeper memory usage is higher than 95 percents
        summary: Some of ZooKeeper Service pod uses memory higher then 95 percents

- interval: 1m
  input_series:
  - series: container_memory_working_set_bytes{namespace="default",container="zookeeper-0",pod="zookeeper-0"}
    values: "0x5"
  - series: kube_pod_container_resource_limits_memory_bytes{exported_namespace="default", exported_pod="zookeeper-0"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Memory_Usage_Alarm
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: zookeeper_backup_metric_last_backup_status{project_name="default"}
    values: "4x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Last_Backup_Has_Failed_Alarm
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: zookeeper-montemplate
        project_name: default
      exp_annotations:
        description: ZooKeeper Last Backup Has Failed
        summary: ZooKeeper Last Backup Has Failed

- interval: 1m
  input_series:
  - series: zookeeper_backup_metric_last_backup_status{project_name="default"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-zookeeper-montemplate
    alertname: ZooKeeper_Last_Backup_Has_Failed_Alarm
    exp_alerts: []